Antes de empezar a resolver los problemas de este notebook, comprobad de que todo funciona como es debido. Primero, **reiniciad el kernel** (En el menú superior, seleccionad Kernel$\rightarrow$Restart) y después **run all cells** (en el menú, seleccionad Cell$\rightarrow$Run All).

Asegúraos también de que rellenais todas las partes que pongan `YOUR CODE HERE` or `YOUR ANSWER HERE`, así como vuestros nombres y apellidos en la siguiente celda. 
Cuando añadais código, aseguraos de eliminar la línea con el `raise NotImplementedError()` para que no dé error.

```python
NOMBRES = ""
```

---

# Tarea 1

## Ejercicio 1 (1.5 puntos)
Escribe una función `numeros_1(limite_inferior, limite_superior)` que devuelva los números que sean divisibles por 7 pero no sean múltiplos de 5 entre dos números dados (ambos incluidos).

<!-- grade_id: cell-8ba2dfd8e9973e4c, points: 0, grade: false, solution: true -->
```python
def numeros_1(limite_inferior, limite_superior):
    l =[] #Aquí se guardan los números que vayamos encontrando
    # Recorremos la lista de números generada con la función range. 
    # Nótese que tenemos que usar limite_superior+1 para generar un intervalo cerrado por la derecha
    for i in range(limite_inferior, limite_superior+1): 
        if i % 7 == 0 and i % 5 != 0: 
            l.append(i)
    return l
```

<!-- grade_id: cell-74c2ff470b100b37, points: 1.5, grade: true, solution: false -->
```python
assert(numeros_1(0, 7) == [7])
assert(numeros_1(0, 14) == [7, 14])
assert(numeros_1(120, 350) == [126,
 133,
 147,
 154,
 161,
 168,
 182,
 189,
 196,
 203,
 217,
 224,
 231,
 238,
 252,
 259,
 266,
 273,
 287,
 294,
 301,
 308,
 322,
 329,
 336,
 343])
```

## Ejercicio 2 (1.5 puntos)

Escribe una función `suma(lista)` que devuelva la suma de todos los elementos de una lista.

<!-- grade_id: cell-dafe590def2b656f, points: 0, grade: false, solution: true -->
```python
def suma(lista):
    #Simplemente usamos un acumulador y el operador += para ir incrementando la cuenta dentro de un bucle for
    acc = 0
    for e in lista:
        acc += e
    return acc
```

<!-- grade_id: cell-0603fcd7515919ea, points: 1.5, grade: true, solution: false -->
```python
assert(suma([1,2,3,4,5]) == 15)
assert(suma([1,2,3,4,5,6,7]) == 28)
```

## Ejercicio 2: Metiéndonos un poco en harina 🥖🍞

La variable cadena que defino debajo alberga el origen de replicación (oriC) del genoma del amigo de John Snow, el *Vibrio Cholerae*:

```python
vc_oric = 'atcaatgatcaacgtaagcttctaagcatgatcaaggtgctcacacagtttatccacaacctgagtggatgacatcaagataggtcgttgtatctccttcctctcgtactctcatgaccacggaaagatgatcaagagaggatgatttcttggccatatcgcaatgaatacttgtgacttgtgcttccaattgacatcttcagcgccatattgcgctggccaaggtgacggagcgggattacgaaagcatgatcatggctgttgttctgtttatcttgttttgactgagacttgttaggatagacggtttttcatcactgactagccaaagccttactctgcctgacatcgaccgtaaattgataatgaatttacatgcttccgcgacgatttacctcttgatcatcgatccgattgaagatcttcaattgttaattctcttgcctcgactcatagccatgatgagctcttgatcatgtttccttaaccctctattttttacggaagaatgatcaagctgctgctcttgatcatcgtttc'
```

Y debajo declaro una cadena que contiene toda la secuencia de ADN del mismo organismo

No te preocupes si no entiendes lo que ocurre: Lo que estoy haciendo es abrir un fichero que contiene la secuencia (más sobre esto en las próximas clases)

```python
vc_sec = open("vibrio_cholerae.txt","r+").read()
#Es una cadena muuuuy larga
#Comprueba su longitud!
vc_sec
```

## Pregunta 1 (1.5 puntos)
¿Qué longitud tienen ambas cadenas? ¿Qué porcentaje del ADN del Vibrio Cholerae representa el origen de la replicación? Imprime las respuestas por pantalla.

<!-- grade_id: cell-afec144bff69d082, points: 1.5, grade: true, solution: true -->
```python
vc_oriC = vc_oriC.upper() 
vc_oriC_len = len(vc_oriC)
vc_sec_len = len(vc_sec) 
print("Longitud oriC: %d" % vc_oriC_len)
print("Longitud secuencia completa: %d" % vc_sec_len)
print("El origen de replicación representa el %.4f%% de la secuencia" % (100*vc_oriC_len/vc_sec_len)) 
```

__Explicación__

Primero paso vc_oriC a mayúsculas por comodidad. Luego Calculo las longitudes usando `len()` y las guardo en sendas variables, que imprimo por pantalla usando `print()`. Finalmente, para calcular el porcentaje, divido la longitud de oriC por la de toda la secuencia y multiplico por 100. 

## Pregunta 2 (2.5 puntos)
¿Cuáles son las frequencia absolutas y relativas de bases {A, C, T, G} en el origen de replicación? ¿Y en todo el genoma? Imprime las respuestas por pantalla.

<!-- grade_id: cell-5e06eedeb257cf7d, points: 2.5, grade: true, solution: true -->
```python
bases = ("A", "C", "T", "G")
for base in bases:
    print("%s: oriC: %d(%.4f%%), genome: %d(%.4f%%)" % (base, 
                                                 vc_oriC.count(base), 
                                                 100*vc_oriC.count(base) / vc_oriC_len,
                                                 vc_sec.count(base),
                                                 100*vc_sec.count(base) / vc_sec_len))
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_880586/2103459827.py in <module>
      2 for base in bases:
      3     print("%s: oriC: %d(%.4f%%), genome: %d(%.4f%%)" % (base, 
----> 4                                                  vc_oriC.count(base),
      5                                                  100*vc_oriC.count(base) / vc_oriC_len,
      6                                                  vc_sec.count(base),

NameError: name 'vc_oriC' is not defined
```

__Explicación__

Itero, usando la expresión `for ... in`, por las bases del alfabeto A,C,T,G que defino como una tupla.
Para cada una, imprimo una sóla línea mostrando los datos que se piden. 
Al imprimirlo de esta manera, es más fácil comparar las frecuencias de los 4 nucleótidos por filas o columnas.
Así, es fácil ver que el porcentaje de timinas es sensiblemente mayor en el origen de replicación que en el resto del ADN.

## Pregunta 3 (3 puntos)
Puedes adaptar el código para que haga lo mismo que en el anterior caso, pero para permutaciones de dos elementos de un alfabeto de 4 nucleótidos `{A, C, G, T}`, es decir, `{AC, AT, AG, CT, CG, TG, CA, TA, GA, TC, GC, GT}`?
Evita definir los elementos de esta serie en una lista o en variables previamente. En su lugar, emplea bucles para generarlos **dinámicamente** en tiempo de ejecución. 

<!-- grade_id: cell-8db935f261e6a91e, points: 3, grade: true, solution: true -->
```python
for i in range(len(bases) - 1):
    for j in range(i+1, len(bases)):
        par = bases[i] + bases[j]
        print("%s: oriC: %d(%.4f%%), genome: %d(%.4f%%)" % (par, 
                                                 vc_oriC.count(par), 
                                                 100*vc_oriC.count(par) / vc_oriC_len,
                                                 vc_sec.count(par),
                                                 100*vc_sec.count(par) / vc_sec_len))
        
        #Si quisiera hacerlo también para las inversas (aunque no se pedía)
        par = bases[j] + bases[i]
        print("%s: oriC: %d(%.4f%%), genome: %d(%.4f%%)" % (par, 
                                                 vc_oriC.count(par), 
                                                 100*vc_oriC.count(par) / vc_oriC_len,
                                                 vc_sec.count(par),
                                                 100*vc_sec.count(par) / vc_sec_len))
```

Output:
```
NameError                                 Traceback (most recent call last)
/tmp/ipykernel_880586/487460667.py in <module>
      3         par = bases[i] + bases[j]
      4         print("%s: oriC: %d(%.4f%%), genome: %d(%.4f%%)" % (par, 
----> 5                                                  vc_oriC.count(par),
      6                                                  100*vc_oriC.count(par) / vc_oriC_len,
      7                                                  vc_sec.count(par),

NameError: name 'vc_oriC' is not defined
```

__Explicación__ 

Ahora se nos pide que hagamos lo mismo pero para dinucléotidos obtenidos como combinaciones sin repetición del alfabeto. 
Para generar los dinucléotidos empleo dos bucles, uno anidado dentro del otro. El más externo se ocupa de generar el primer nucléotido, mientras que el interno genera el segundo. 

Como nos piden permutaciones sin repetición, no se repetirán nucléotidos (AA, GG, AAT, ATT, etc. no serían válidas). Una manera programática de generar estas combinaciones es iterar por cada elemento de la lista (bucle externo) y combinarlo con todas aquellas bases que se encuentren a su derecha (bucle interno).

Una vez obtenido el dinucleótido, puedo calcular las frecuencias de manera análoga al ejercicio anterior. Si quisiera analizar las inversas, bastaría con permutar el orden de los elementos (índices i y j). 

Es importante reseñar que este enfoque no es muy escalable, ya que el número de bucles anidados que hay que usar para generar los resultados deseados depende de la longitud de las combinaciones 

```python

```